---
title: 'Insert the chapter title here'
description: 'Insert the chapter description here'
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/production/course_6215/slides/chapter1.pdf'
free_preview: true
---

## changed while unsaved changes

```yaml
type: NormalExercise
key: 6134f7e696
xp: 100
```

changeda
aaacommitaa

`@instructions`
Please enter testall the data you need bla bla bla

`@hint`
Agree with their instructor to use a new one or old one

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}
dsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkals
```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 7686998169
xp: 50
```

How much is too much? `boom` aaa

`@hint`
A lota

`@possible_answers`
- 1
- 2
- [3]
- 4

`@feedback`
- not ok
- not ok
- bravo!
- not ok

---

## MultipleChoice

```yaml
type: MultipleChoiceExercise
key: c5b95e9c77
xp: 50
```

assignment

`@possible_answers`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```

---

## Tab

```yaml
type: TabExercise
key: ff2903962d
xp: 100
```

this is a tab exercise

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

***

```yaml
type: NormalExercise
key: f289e04658
xp: 25
```

`@instructions`
instructions 1

`@hint`


`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

***

```yaml
type: MultipleChoiceExercise
key: 087c71b069
xp: 25
```

`@question`


`@sample_code`
```{r}

```

`@possible_answers`


`@hint`


`@sct`
```{r}

```

***

```yaml
type: PureMultipleChoiceExercise
key: 19b0a766ff
xp: 25
```

***

```yaml
type: NormalExercise
key: b3503f176d
xp: 25
```

`@instructions`


`@hint`


`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Pure Multiple Choice

```yaml
type: PureMultipleChoiceExercise
key: f260bef449
xp: 50
```

how are you today

`@hint`
Im fine thanks

`@possible_answers`


`@feedback`
- good
- bad

---

## Bullet

```yaml
type: BulletExercise
key: b4986fd506
xp: 100
```

this is the assignment and it seems to be working. Goranhere testing HotJar.

`@pre_exercise_code`
```{r}
code
```

`@sample_code`
```{r}
code
```

***

```yaml
type: NormalExercise
key: f05b7df4d4
xp: 25
```

`@instructions`


`@hint`


`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

***

```yaml
type: MultipleChoiceExercise
key: c9e5d11a54
xp: 25
```

`@question`


`@sample_code`
```{r}

```

`@possible_answers`


`@hint`


`@sct`
```{r}

```

***

```yaml
type: PureMultipleChoiceExercise
key: 31d077660f
xp: 25
```

***

```yaml
type: MultipleChoiceExercise
key: 205e240d34
xp: 25
```

`@question`


`@sample_code`
```{r}

```

`@possible_answers`
ssa

`@hint`
aa

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: e8aa0cd19d
xp: 50
```

dasdsa

`@hint`
dsadas

`@possible_answers`


`@feedback`
- yes
- no

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 9833955dec
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 34ce4ccf10
xp: 50
```

mac

`@possible_answers`
- 1
- 2
- 3

`@hint`
das

`@pre_exercise_code`
```{r}
# pre ec
```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: f8eb34a1b6
xp: 50
```

assignment

`@hint`
ok seems to be workign

`@possible_answers`


`@feedback`
- correct
- not correct

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: c82cf1f1b8
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: c5a0dda76b
xp: 50
```

`@projector_key`
9b8e5a8ffff6b5070b7ec4fd1193329b

---

## Insert exercise title here

```yaml
type: NormalExercise
key: e0ccadecd7
xp: 100
```

change

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: ebfabd1841
xp: 50
```

`@projector_key`
f78b529cd0840764af00ebc2876f751d

---

## <<<New Exercise>>>

```yaml
type: MultipleChoiceExercise
key: fc9e67768c
lang: r
xp: 50
skills: 1
```

this is a question

`@possible_answers`
- instruction 1
- instruction 2

`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 55aa9d936f
xp: 50
```

`@projector_key`
08521096d1ea8d0b3b8102175b708690

---

## <<<New Exercise>>>

```yaml
type: PureMultipleChoiceExercise
key: 9be5906e74
xp: 50
skills: 1
```



`@hint`


`@possible_answers`


`@feedback`
- good
- bad

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 69ee6b12cb
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: f9c9c6887a
xp: 50
```

`@projector_key`
d8b0618ad679aa179216c5cbd8616ae7

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 40d4510f1f
xp: 50
```

`@projector_key`
12fc66a2e114033bef2e87e038961af6

---

## Insert exercise title here

```yaml
type: NormalExercise
key: ba34234bd3
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
