---
title: 'Insert the chapter title here'
description: 'Insert the chapter description here'
---

## testing long long testing long long testing long long testing long long testing long long testing long long testing long long testing long long sada

```yaml
type: NormalExercise
key: 6134f7e696
xp: 100
```

changeda

`@instructions`


`@hint`
Agree with their instructor to use a new one or old one

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}
dsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkalsdsanlkdnaskldnkals
```

`@possible_answers`


`@feedback`


---

## MultipleChoice

```yaml
type: MultipleChoiceExercise
key: c5b95e9c77
xp: 50
```

assignment

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


---

## Pure Multiple Choice

```yaml
type: PureMultipleChoiceExercise
key: f260bef449
xp: 50
```

how are you today

`@instructions`


`@hint`
Im fine thanks

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`
- [i'm ok]
- not ok

`@feedback`
- good
- bad

---

## Tab

```yaml
type: TabExercise
key: ff2903962d
xp: 100
```

this is a tab exercise

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: f289e04658
xp: 25
```



`@instructions`
instructions 1

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 087c71b069
xp: 25
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

***

```yaml
type: PureMultipleChoiceExercise
key: 19b0a766ff
xp: 25
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

***

```yaml
type: NormalExercise
key: b3503f176d
xp: 25
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

---

## Bullet

```yaml
type: BulletExercise
key: b4986fd506
xp: 100
```

this is the assignment and it seems to be working. Goranhere testing HotJar.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
code
```

`@sample_code`
```{r}
code
```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: f05b7df4d4
xp: 25
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: c9e5d11a54
xp: 25
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

***

```yaml
type: PureMultipleChoiceExercise
key: 31d077660f
xp: 25
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 205e240d34
xp: 25
```



`@instructions`
ssa

`@hint`
aa

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: e8aa0cd19d
xp: 50
```

dasdsa

`@instructions`


`@hint`
dsadas

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`
- [ok]
- not ok

`@feedback`
- yes
- no

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 9833955dec
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 34ce4ccf10
xp: 50
```

mac

`@instructions`
- 1
- 2
- 3

`@hint`
das

`@pre_exercise_code`
```{r}
# pre ec
```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: f8eb34a1b6
xp: 50
```

assignment

`@instructions`


`@hint`
ok seems to be workign

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`
- [1]
- 2

`@feedback`
- correct
- not correct

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: c82cf1f1b8
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: c5a0dda76b
xp: 50
```

`@projector_key`
9b8e5a8ffff6b5070b7ec4fd1193329b

---

## Insert exercise title here

```yaml
type: NormalExercise
key: e0ccadecd7
xp: 100
```

change

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: ebfabd1841
xp: 50
```

`@projector_key`
f78b529cd0840764af00ebc2876f751d

---

## <<<New Exercise>>>

```yaml
type: MultipleChoiceExercise
key: fc9e67768c
lang: r
xp: 50
skills: 1
```

this is a question

`@instructions`
- instruction 1
- instruction 2

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: 55aa9d936f
xp: 50
```

`@projector_key`
08521096d1ea8d0b3b8102175b708690

---

## <<<New Exercise>>>

```yaml
type: PureMultipleChoiceExercise
key: 9be5906e74
xp: 50
skills: 1
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`
- [1]
- 2

`@feedback`
- good
- bad

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 69ee6b12cb
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`
