---
title: '4'
description: ""
---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: 6884254136   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


***



```yaml
type: NormalExercise 
xp: 50 
key: 78c675bafb   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`



***



```yaml
type: NormalExercise 
xp: 50 
key: 5bcd93a328   
```





`@instructions`
assignment

`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


