---
title: second
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: f7dab399fa
xp: 100
```

goran

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 9f265fe16c
xp: 50
```

`@projector_key`
a6ffd10835703a3371821c0e25ef1ee7

---

## Insert exercise title here

```yaml
type: VideoExercise
key: d4e3dc890d
xp: 50
```

`@projector_key`
6fcf301e49613465d27d4b24b6af72e6

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 06723000f5
xp: 50
```

`@projector_key`
19fe255d261c6d7c99ce35032c096856

---

## Insert exercise title here

```yaml
type: VideoExercise
key: ff82ced04d
xp: 50
```

`@projector_key`
cd1ec04722fccb740c1a75984e53497b

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 3511ef9d9c
xp: 50
```

`@projector_key`
64d1d77fd02ccea64155854db2f35aed
