---
title: dasdas
description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise
key: de65da2c7b
xp: 50
```

`@projector_key`
99c3205a70eb25cc53db08825888a2c2

---

## Insert exercise title here

```yaml
type: TabExercise
key: bd12611c31
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: e4722dd49e
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: fd2159bdcd
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`

---

## Insert exercise title here

```yaml
type: BulletExercise
key: 6333029132
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: 4f0b9a52c1
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

`@possible_answers`


`@feedback`
