---
title: '4'
description: ""
---

## Insert exercise title here

```yaml
type: TabExercise
key: 6884254136
xp: 100
```



`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

***

```yaml
type: NormalExercise
key: 78c675bafb
xp: 50
```



`@instructions`


`@hint`


`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

***

```yaml
type: NormalExercise
key: 5bcd93a328
xp: 50
```



`@instructions`
assignment

`@hint`


`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
