---
title: 'add chapter'
description: ""
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/production/course_6215/slides/chapter6.pdf'
---

## Insert exercise title here

```yaml
type: VideoExercise
key: 6d078f6ea4
xp: 50
```

`@projector_key`
6be09036aa1d80a38b590d1fc8290ce9

---

## goran

```yaml
type: NormalExercise
key: e63ff47484
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
