---
  title: "Insert the chapter title here"
  description: "Insert the chapter description here"
  v2: true

---
## Normal

```yaml
type: NormalExercise 
xp: 100 
key: 6134f7e696   
```

changed



`@pre_exercise_code`

```{undefined}
daskdnaskldnklasndklas
daslkdnaskldnklas
dasdalkndakslndklasnkldankld
d
a
da
d
a
da
```









---
## MultipleChoice

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: c5b95e9c77   
```

assignment













---
## Pure Multiple Choice

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: f260bef449   
```

how are you today


`@hint`
Im fine thanks





`@possible_answers`
- [i'm ok]
- not ok

`@feedbacks`
- good
- bad





---
## Tab

```yaml
type: TabExercise 
xp: 100 
key: ff2903962d   
```

this is a tab exercise












***



```yaml
type: NormalExercise 
xp: 35 
key: f289e04658   
```



`@instructions`
instructions 1

`@hint`













***



```yaml
type: MultipleChoiceExercise 
xp: 35 
key: 087c71b069   
```



`@instructions`


`@hint`













***



```yaml
type: PureMultipleChoiceExercise 
xp: 30 
key: 19b0a766ff   
```




`@hint`






`@possible_answers`


`@feedbacks`







---
## Bullet

```yaml
type: BulletExercise 
xp: 100 
key: b4986fd506   
```

this is the assignment and it seems to be working. Goranhere testing HotJar.



`@pre_exercise_code`

```{undefined}
code
```
`@sample_code`

```{undefined}
code
```







***



```yaml
type: NormalExercise 
xp: 25 
key: f05b7df4d4   
```



`@instructions`


`@hint`













***



```yaml
type: MultipleChoiceExercise 
xp: 25 
key: c9e5d11a54   
```



`@instructions`


`@hint`













***



```yaml
type: PureMultipleChoiceExercise 
xp: 25 
key: 31d077660f   
```




`@hint`






`@possible_answers`


`@feedbacks`







***



```yaml
type: MultipleChoiceExercise 
xp: 25 
key: 205e240d34   
```



`@instructions`
ssa

`@hint`
aa












---
## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: e8aa0cd19d   
```

dasdsa


`@hint`
dsadas





`@possible_answers`
- [ok]
- not ok

`@feedbacks`
- yes
- no





---
## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 9833955dec   
```















---
## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: f8eb34a1b6   
```

assignment


`@hint`
ok seems to be workign





`@possible_answers`
- [1]
- 2

`@feedbacks`
- correct
- not correct





---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 34ce4ccf10   
```

mac

`@instructions`
- 1
- 2
- 3

`@hint`
das

`@pre_exercise_code`

```{undefined}
# pre ec
```









---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: c82cf1f1b8   
```















---
## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: c5a0dda76b   
```

`@projector_key`
9b8e5a8ffff6b5070b7ec4fd1193329b

---
## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: e0ccadecd7   
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{undefined}

```

`@sample_code`

```{undefined}

```

`@solution`

```{undefined}

```

`@sct`

```{undefined}

```





