---
title: second
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: f7dab399fa   
```


goran


`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 9f265fe16c   
```

`@projector_key`
a6ffd10835703a3371821c0e25ef1ee7
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: d4e3dc890d   
```

`@projector_key`
6fcf301e49613465d27d4b24b6af72e6
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 06723000f5   
```

`@projector_key`
19fe255d261c6d7c99ce35032c096856
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: ff82ced04d   
```

`@projector_key`
cd1ec04722fccb740c1a75984e53497b
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 3511ef9d9c   
```

`@projector_key`
64d1d77fd02ccea64155854db2f35aed