---
title: dasdas
description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: de65da2c7b   
```

`@projector_key`
99c3205a70eb25cc53db08825888a2c2
---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: bd12611c31   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


***



```yaml
type: NormalExercise 
xp: 50 
key: e4722dd49e   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`



***



```yaml
type: MultipleChoiceExercise 
xp: 50 
key: fd2159bdcd   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`



---

## Insert exercise title here

```yaml
type: BulletExercise 
xp: 100 
key: 6333029132   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


***



```yaml
type: NormalExercise 
xp: 100 
key: 4f0b9a52c1   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


`@possible_answers`


`@feedback`


