---
  title: "4"
  description: ""
---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: 6884254136   
```





***



```yaml
type: NormalExercise 
xp: 50 
key: 78c675bafb   
```





`@instructions`


`@hint`


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```



***



```yaml
type: NormalExercise 
xp: 50 
key: 5bcd93a328   
```





`@instructions`
assignment

`@hint`


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


